A common alteration in effort-based decision-making in apathy, anhedonia, and late circadian rhythm

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

This study uses an online cognitive task to assess how reward and effort are integrated in a motivated decision-making task. In particular the authors were looking to explore how neuropsychiatric symptoms, in particular, apathy and anhedonia, and circadian rhythms affect behavior in this task. Amongst many results, they found that choice bias (the degree to which integrated reward and effort affect decisions) is reduced in individuals with greater neuropsychiatric symptoms, and late chronotypes (being an 'evening person').

Strengths:

The authors recruited participants to perform the cognitive task both in and out of sync with their chronotypes, allowing for the important insight that individuals with late chronotypes show a more reduced choice bias when tested in the morning.

Overall, this is a well-designed and controlled online experimental study. The modelling approach is robust, with care being taken to both perform and explain to the readers the various tests used to ensure the models allow the authors to sufficiently test their hypotheses.

Weaknesses:

This study was not designed to test the interactions of neuropsychiatric symptoms and chronotypes on decision making, and thus can only make preliminary suggestions regarding how symptoms, chronotypes and time-of-assessment interact.

Reviewer #2 (Public Review):

Summary:

The study combines computational modeling of choice behavior with an economic, effort-based decision-making task to assess how willingness to exert physical effort for a reward varies as a function of individual differences in apathy and anhedonia, or depression, as well as chronotype. They find an overall reduction in effort selection that scales with apathy, anhedonia and depression. They also find that later chronotypes are less likely to choose effort than earlier chronotypes and, interestingly, an interaction whereby later chronotypes are especially unwilling to exert effort in the morning versus the evening.

Strengths:

This study uses state-of-the-art tools for model fitting and validation and regression methods which rule out multicollinearity among symptom measures and Bayesian methods which estimate effects and uncertainty about those estimates. The replication of results across two different kinds of samples is another strength. Finally, the study provides new information about the effects not only of chronotype but also chronotype by timepoint interactions which are previously unknown in the subfield of effort-based decision-making.

Weaknesses:

The study has few weaknesses. The biggest drawback is that it does not provide evidence for the idea that a match between chronotype and delay matters is especially relevant for people with depression or continuous measures like anhedonia and apathy. It is unclear whether disorders further interact with chronotype and time of day to determine a bias against effort. On the other hand, the study does provide evidence that future studies should consider such interactions when examining questions about effort expenditure in psychiatric disorders.

Reviewer #3 (Public Review):

Summary:

In this manuscript, Mehrhof and Nord study a large dataset of participants collected online (n=958 after exclusions) who performed a simple effort-based choice task. They report that the level of effort and reward influence choices in a way that is expected from prior work. They then relate choice preferences to neuropsychiatric syndromes and, in a smaller sample (n<200), to people's circadian preferences, i.e., whether they are a morning-preferring or evening-preferring chronotype. They find relationships between the choice bias (a model parameter capturing the likelihood to accept effort-reward challenges, like an intercept) and anhedonia and apathy, as well as chronotype. People with higher anhedonia and apathy and an evening chronotype are less likely to accept challenges (more negative choice bias). People with an evening chronotype are also more reward sensitive and more likely to accept challenges in the evening, compared to the morning.

Strengths:

This is an interesting and well-written manuscript which replicates some known results and introduces a new consideration related to chronotype relationships which have not been explored before. It uses a large sample size and includes analyses related to transdiagnostic as well as diagnostic criteria.

Weaknesses:

The authors do not explore how chronotype and depression are related (does one mediate the effect of the other etc). Both variables are included in the same model in the revised article now which is a great improvement, but it also means psychopathology and circadian rhythms are treated as distinct phenomena and their relationship in predicting effort-reward preferences is not examined.

Recommendations for the authors:

Reviewer #3 (Recommendations For The Authors):

Two points in response to changes the authors made:

(1) "motivational tendency" is in our opinion not an improved phrase over "choice bias". A paper by Jon Roiser calls it "overall bias to accept effortful challenges" (but that's maybe too long?)

We thank the reviewer for their suggestion of renaming our computational parameter and agree it would be of value to introduce and label this parameter in line with other work, improving consistency across the literature. Hence, we have updated our manuscript and now introduce the parameter as bias to accept effortful challenges for reward and refer to the parameter as acceptance bias thereafter.

We have updated this nomenclature throughout the manuscript text, figures and supplement.

(2) The new title "Both neuropsychiatric symptoms and circadian rhythm alter effort-based decision-making" sounds slightly causal (as would be the case in a longitudinal or intervention study). Maybe instead the authors could use "are associated with" or similar?

We agree with the reviewers that our current title could be interpreted in a causal manner. We have updated our title to now read A common alteration in effort-based decision-making in apathy, anhedonia, and late circadian rhythm.

Changes from the earlier version to the revised version

Abstract

Abstract

Motivational deficits are common in several brain disorders, and motivational syndromes like apathy and anhedonia predict worse outcomes. Disrupted effort-based decision-making may represent a neurobiological underpinning of motivational deficits, shared across neuropsychiatric disorders. We measured effort-based decision-making in 994 participants using a gamified online task, combined with computational modelling, and validated offline for test-retest reliability. In two pre-registered studies, we first replicated studies linking impaired effort-based decision-making to neuropsychiatric syndromes, taking both a transdiagnostic and a diagnostic-criteria approach. Next, testing participants with early and late circadian rhythms in the morning and evening, we find circadian rhythm interacts with time-of-testing to produce parallel effects on effort-based decision-making. Circadian rhythm may be an important variable in computational psychiatry, decreasing reliability or distorting results when left unaccounted for. Disentangling effects of neuropsychiatric syndromes and circadian rhythm on effort-based decision-making will be essential to understand motivational pathologies and to develop tailored clinical interventions.

1 Introduction

Our circadian rhythm aligns us with our environment, regulating physiological and behavioural processes to follow 24-hour rhythms1. Circadian integrity is pivotal to mental wellbeing and has been bidirectionally linked to numerous psychiatric disorders2–8. Yet little is known about the cognitive or computational mechanisms of circadian dysfunction—and their alignment or diversion from mechanisms driving neuropsychiatric symptoms.

Earlier version: Inter-individual differences in circadian timing and alignment manifest behaviourally as chronotypes (i.e., diurnal preference)9,10, with individuals commonly categorized as early, late, or intermediate chronotypes11. A disproportionate number of psychiatric patients have a late chronotype, based on self-report12 and genetic analysis9. Within clinical groups, late chronotype has been linked to depression severity and non-remission13, higher rates of psychiatric and general medical comorbidities14, more severe cognitive impairment, and higher symptoms of apathy15,16. Converging evidence on the importance of circadian alignment in psychiatric pathology has led to proposals of a circadian psychiatric phenotype, either within disorders14,17 or cuding across diagnostic categories18,19.

Revised version: Inter-individual differences in circadian timing and alignment manifest behaviourally as chronotypes (i.e., diurnal preference)9,10, with individuals commonly categorized as early, late, or intermediate chronotypes11. A disproportionate number of psychiatric patients have a late chronotype, based on self-report12 and genetic analysis9. Within clinical groups, late chronotype has been linked to depression severity and non-remission13, higher rates of psychiatric and general medical comorbidities14, more severe cognitive impairment, and higher symptoms of apathy15,16. Converging evidence on the importance of circadian alignment in psychiatric pathology has led to proposals of a circadian psychiatric phenotype, either within disorders14,17 or cutting across diagnostic categories18,19.

Syndromes of deficient motivational behaviour, such as apathy and anhedonia, are also observed across neuropsychiatric disorders20–23, suggesting transdiagnostic relevance24. Anhedonia and apathy are associated with worse clinical outcomes25,26 and are poorly targeted by current treatments27–29. Empirical work suggests a common underlying neurocognitive mechanism: the integration of costs and benefits during effortful decision-making24. Effort-based decision-making is commonly assessed using effort expenditure tasks: Subjects are asked to decide whether to pursue actions associated with varying levels of effort and reward levels30. Computational models applied to effort-based decision-making tasks provide a formal mathematical estimate of a subject’s integration of costs and benefits into a subjective value31,32. Higher costs devalue associated rewards, an effect referred to as effort-discounting33–37. This computational approach enables measurement of inter- and intra-individual differences on distinct aspects of effort-based decision-making.

One key source of individual differences in motivational behaviour and effort-based decision-making is likely dopamine signalling, especially dopaminergic projections from the ventral tegmental area (VTA) to the ventral striatum24. Pre-clinical animal studies show dopamine depletion reduces engagement in effortful behaviour38,39, while dopamine enhancement promotes motivational effort exertion40,41. In humans, dopamine depletion reduces willingness to exert effort for reward42,43, while pharmacological dopamine enhancement increases motivation in effort-based decision-making44,45. Further, naturally occurring variations in dopamine responsivity are correlated with effort-based decision-making: a higher dopamine responsivity (as quantified with positron emission tomography following d-amphetamine administration) is associated with willingness to exert greater effort for larger rewards46.

Bi-directional links between chronobiology and several neurotransmitter systems have been reported, including dopamine47. In animals, dopamine transmission and biosynthesis vary diurnally48,49, and growing evidence suggests a bi-directional regulation between dopamine signalling and circadian rhythm50–52. In human studies, dopamine availability, dopamine transporter genes, and dopamine receptors have been linked to proxies of circadian rhythm53–55 and circadian-regulating gene polymorphisms56. On a behavioural level, sleep deprivation, poor sleep quality, and insomnia were linked to low motivation in effort-based decision-making57–59 and evening bright-light exposure enhanced effort willingness, possibly by enhancing dopamine through melatonin suppression60. Early chronotype predicted treatment effect on motivational behaviour in a sample of depressed subjects with comorbid insomnia61. Chronotype effects are also reported for other reward decision-making tasks, with late chronotypes showing higher delay discounting62, less rational decision-making63, and lower willingness to take risks for rewards64,65. A circadian effect on decision-making under risk is reported, with the sensitivity to losses decreasing with time-of-day66. This suggests that chronobiology may contribute to individual differences in effort-based decision-making, potentially in parallel ways with neuropsychiatric syndromes.

Here, we tested the relationship between motivational decision-making and three key neuropsychiatric syndromes: anhedonia, apathy, and depression, taking both a transdiagnostic and categorical (diagnostic) approach. To do this, we validate a newly developed effort-expenditure task, designed for online testing, and gamified to increase engagement. Participants completed the effort-expenditure task online, followed by a series of self-report questionnaires.

Earlier version: Next, we pre-registered a follow-up experiment to directly investigate how circadian preference interacts with time-of-day on motivational decision-making, using the same task and computational modelling approach. While this allows us to test how circadian effects on motivational decision-making compare to neuropsychiatric effects, we do not test for possible interactions between neuropsychiatric symptoms and chronobiology. All analyses were pre-registered (except when labelled as exploratory): see hdps://osf.io/2x3au and hdps://osf.io/y4ke.

Revised version: Next, we pre-registered a follow-up experiment to directly investigate how circadian preference interacts with time-of-day on motivational decision-making, using the same task and computational modelling approach. While this allows us to test how circadian effects on motivational decision-making compare to neuropsychiatric effects, we do not test for possible interactions between neuropsychiatric symptoms and chronobiology. All analyses were pre-registered (except when labelled as exploratory): see https://osf.io/2x3au and https://osf.io/y4fbe.

2 Results

2.1 Sample characteristics

Nine hundred and ninety-four participants completed all study components (i.e., demographic questions, effort-expenditure task, self-report questionnaires). After exclusion (see Methods 4.1.5), 958 participants were included in our analyses. We used a stratified recruitment approach to ensure our sample was representative of the UK population in age, sex, and history of psychiatric disorder67–69; mean questionnaire-based measures were comparable to previous general population studies (Table 1).

Correlations between questionnaire scores:

Correlations between questionnaire sum scores for the Snaith Hamilton Pleasure Scale (SHAPS), the Dimensional Anhedonia Rating Scale (DARS), the Apathy Evaluation Scale (AES), Morningness-Eveningness Questionnaire (MEQ), Munich Chronotype Questionnaire (MCTQ), Body mass index (BMI) and the Finish Diabetes Risk Score (FINDRISC). Asterisks indicate significance: * p<.05, ** p<.01. *** p<.001 (not accounting for multiple comparisons). Note that sum scores for the AES and the DARS have been transformed such that increasing scores can be interpreted as higher symptom severity, in line with the SHAPS. Sum scores of the MEQ have been transformed such that higher scores indicate higher eveningness, in line with the MCTQ.

Questionnaire sum scores highly correlated within groupings of questionnaires targeting psychiatric symptoms, chronobiology, and metabolic health. We also found significant correlations between some, but not all, questionnaires (Fig. 1)

2.2 Effort-expenditure task

Earlier version: In this novel, online effort-expenditure task (Fig. 2A-B), subjects were given a series of challenges associated with varying levels of effort and reward. By weighing up efforts against rewards, they decide whether to accept or reject challenges. We first use model-agnostic analyses to replicate effects of effort-discounting (i.e., devaluation of reward with increasing effort). Next, we took a computational modelling approach to fit economic decision-making models to the task data (Fig.3A-D). The models posit efforts and rewards are joined into a subjective value (SV), weighed by individual effort (βE) and reward sensitivity (βR) parameters. The subjective value is then integrated with an individual motivational tendency (a) parameter to guide decision-making. Specifically, the motivational tendency parameter determines the range at which subjective values are translated to acceptance probabilities: the same subjective value will translate to a higher acceptance probability the higher the motivational tendency.

Revised version: In this novel, online effort-expenditure task (Fig. 2A-B), subjects were given a series of challenges associated with varying levels of effort and reward. By weighing up efforts against rewards, they decide whether to accept or reject challenges. We first use model-agnostic analyses to replicate effects of effort-discounting (i.e., devaluation of reward with increasing effort). Next, we took a computational modelling approach to fit economic decision-making models to the task data (Fig. 3A-D). The models posit efforts and rewards are joined into a subjective value (SV), weighed by individual effort (βE) and reward sensitivity (βR) parameters. The subjective value is then integrated with an individual bias to accept effortful challenges for reward (α) parameter to guide decision-making. Specifically, this acceptance bias parameter determines the range at which subjective values are translated to acceptance probabilities: the same subjective value will translate to a higher acceptance probability the higher the acceptance bias.

Effort-based decision-making: task design and model-agnostic results.

A: The task can be divided into four phases: a calibration phase to determine individual clicking capacity to calibrate effort-levels, practice trials that participants practice until successful on every effort-level, instructions and a quiz that must be passed, and the main task, consisting of 64 trials split into 4 blocks. B: Each trial consists of an offer with a reward (2,3,4, or 5 points) and an effort level (1,2,3, or 4, scaled to the required clicking speed and time the clicking must be sustained for) that subjects accept or reject. If accepted, a challenge at the respective effort level must be fulfilled for the required time to win the points. If rejected, subjects wait for a matched amount of time and receive one point. C: Proportion of accepted trials, averaged across participants and effort-reward combinations. Error bars indicate standard errors. D: Staircasing development of offered effort and reward levels across the task, averaged across participants.

Computational modelling: model visualization and model-based results.

Earlier version: A: Economic decision-making models posit that efforts and rewards are joined into a subjective value (SV), weighed by individual effort (βE) and reward sensitivity (βR) parameters. The SV is then integrated with a motivational tendency parameter and translated to decision-making. B-C: The model suggests that SV decreases as effort increases and increases as reward increases. The magnitude of this relationship depends on the individual effort and reward sensitivity parameters. D: The motivational tendency parameter acts as an intercept to the softmax function, thereby changing the relationship between SV and acceptance probability. E: Model comparison based on leave-out-out information criterion (LOOIC; lower is beger) and expected log posterior density (ELPD; higher is beger). F: Posterior predictive checks for the full parabolic model, comparing observed vs. model predicted subject-wise acceptance proportions across effort-levels (left) and reward-levels (right).

Revised version: A: Economic decision-making models posit that efforts and rewards are joined into a subjective value (SV), weighed by individual effort (βE) and reward sensitivity (βR) parameters. The SV is then integrated with an acceptance bias parameter and translated to decision-making. B-C: The model suggests that SV decreases as effort increases and increases as reward increases. The magnitude of this relationship depends on the individual effort and reward sensitivity parameters. D: The acceptance bias parameter acts as an intercept to the softmax function, thereby changing the relationship between SV and acceptance probability. E: Model comparison based on leave-out-out information criterion (LOOIC; lower is better) and expected log posterior density (ELPD; higher is better). F: Posterior predictive checks for the full parabolic model, comparing observed vs. model predicted subject-wise acceptance proportions across effort-levels (left) and reward-levels (right).

2.2.1 Replication of model-agnostic effects

Earlier version: The proportion of accepted trials for each effort-reward combination is ploded in Figure 2C. In line with our pre-registered hypotheses, we found significant main effects for effort (F(1,14367)=4961.07, p<.0001) and reward (F(1,14367)=3037.91, p<.001), and a significant interaction between the two (F(1,14367)=1703.24, p<.001). In post hoc ANOVAs, effort effects remained significant at all reward-levels (all p<.001) and reward effects remained significant at all effort-levels (all p<.001). The development of offered effort and reward levels across trials is shown in Figure 2D; this shows that as participants generally tend to accept challenges rather than reject them, the implemented staircasing procedure develops toward higher effort and lover reward challenges.

Revised version: The proportion of accepted trials for each effort-reward combination is plotted in Figure 2C. In line with our pre-registered hypotheses, we found significant main effects for effort (F(1,14367)=4961.07, p<.0001) and reward (F(1,14367)=3037.91, p<.001), and a significant interaction between the two (F(1,14367)=1703.24, p<.001). In post hoc ANOVAs, effort effects remained significant at all reward-levels (all p<.001) and reward effects remained significant at all effort-levels (all p<.001). The development of offered effort and reward levels across trials is shown in Figure 2D; this shows that as participants generally tend to accept challenges rather than reject them, the implemented staircasing procedure develops toward higher effort and lover reward challenges.

Earlier version: The mean success rate of accepted challenges across participants was high (M=98.7%) and varied lidle between participants (SD=3.50), indicating feasibility of all effort-levels across participants. Comparing clicking calibration results from pre-to post-task, the maximum clicking capacity decreased by 2.34 clicks on average (SD=14.5). Sixty-two (6.47%) participants reported having deviated from our instructions (i.e., changed the hand and/or finger used to make mouse clicks) throughout the game, but all effects could still be replicated in this subsample: main and interaction effects of effort and reward on the proportion of accepted trials could be replicated in this subsample (all p<.001) and there was no significant difference between participants that did or did not report finger switching in the mean percentage of accepted trials (switching: 79.51%, no switching: 76.62%; p=.149).

Revised version: The mean success rate of accepted challenges across participants was high (M=98.7%) and varied little between participants (SD=3.50), indicating feasibility of all effort-levels across participants. Comparing clicking calibration results from pre- to post-task, the maximum clicking capacity decreased by 2.34 clicks on average (SD=14.5). Sixty-two (6.47%) participants reported having deviated from our instructions (i.e., changed the hand and/or finger used to make mouse clicks) throughout the game, but all effects could still be replicated in this subsample: main and interaction effects of effort and reward on the proportion of accepted trials could be replicated in this subsample (all p<.001) and there was no significant difference between participants that did or did not report finger switching in the mean percentage of accepted trials (switching: 79.51%, no switching: 76.62%; p=.149).

Subjects were engaged with the task, shown by a high rate of challenge acceptance (M=76.80%, SD=15.20, range=15.60–100%) and moderate-to-good enjoyment ratings (M=2.56, SD=0.92; on a 0–4 scale). Qualitative data of subjects describing their decision-making process during the task further confirmed high engagement (see Supplement 3).

2.2.2 Computational modelling

Earlier version: A model space of nine models was considered, varying in the implemented parameter and cost function (see Supplement 1.1 for mathematical definitions of all models). Prior to model fiding, parameter recovery confirmed all models yield meaningful parameter estimates (Supplement 1.2). All models showed good convergence (effective sample size (ESS)>4,223; R-hats<1.002 for all estimates). Model comparison by out-of-sample predictive accuracy identified the model implementing three parameters (motivational tendency α, reward sensitivity βR, and effort sensitivity βE), with a parabolic cost function (subsequently referred to as the full parabolic model) as the winning model (leave-one-out information criterion [LOOIC; lower is beder] = 29734.8; expected log posterior density [ELPD; higher is beder] = - 14867.4; Fig. 3E). This was in line with our pre-registered hypotheses. Predictive validity of the full parabolic model was validated with posterior predictive checks, showing excellent accordance between observed and model-predicted choice data (across effort-levels: R2=.95, across reward-levels: R2=.94; Fig.3F).

Revised version: A model space of nine models was considered, varying in the implemented parameter and cost function (see Supplement 1.1 for mathematical definitions of all models). Prior to model fitting, parameter recovery confirmed all models yield meaningful parameter estimates (Supplement 1.2). All models showed good convergence (effective sample size (ESS)>4,223; R-hats<1.002 for all estimates). Model comparison by out-of-sample predictive accuracy identified the model implementing three parameters (acceptance bias α, reward sensitivity βR, and effort sensitivity βE), with a parabolic cost function (subsequently referred to as the full parabolic model) as the winning model (leave-one-out information criterion [LOOIC; lower is better] = 29734.8; expected log posterior density [ELPD; higher is better] = -14867.4; Fig. 3E). This was in line with our pre-registered hypotheses. Predictive validity of the full parabolic model was validated with posterior predictive checks, showing excellent accordance between observed and model-predicted choice data (across effort-levels: R2=.95, across reward-levels: R2=.94; Fig.3F).

2.2.3 Test-retest reliability

Earlier version: We validated the task in a smaller in-person sample (N=30, tested twice ~7 days apart, holding time-of-day at testing constant) to assess test-retest reliability of parameter estimates, showing moderate to excellent reliability for all parameters (i.e., all intraclass correlation coefficients >0.4, all p<.01). Parameter estimates from modelling the data at one session predicted subjects’ choices at the other session beder than chance and beder than group-level parameters predictions (all p<0.01)72 (full details reported in Supplement 2).

Revised version: We validated the task in a smaller in-person sample (N=30, tested twice ~7 days apart, holding time-of-day at testing constant) to assess test-retest reliability of parameter estimates, showing moderate to excellent reliability for all parameters (i.e., all intraclass correlation coefficients >0.4, all p<.01). Parameter estimates from modelling the data at one session predicted subjects’ choices at the other session better than chance and better than group-level parameters predictions (all p<0.01)72 (full details reported in Supplement 2).

2.3 Transdiagnostic analysis: Questionnaire measures predict effort-based decision-making

Earlier version: We used partial-least-squares (PLS) regression to relate individual-level mean posterior parameter values resulting from the model fiding of the full parabolic model to the questionnaire measures. To explore individual effects post-hoc, we followed up on effects found in the PLS regression using Bayesian generalised linear models (GLMs), controlling for age and gender.

Revised version: We used partial-least-squares (PLS) regression to relate individual-level mean posterior parameter values resulting from the model fitting of the full parabolic model to the questionnaire measures. To explore individual effects post-hoc, we followed up on effects found in the PLS regression using Bayesian generalised linear models (GLMs), controlling for age and gender.

Earlier version: 2.3.1 Motivational tendency

Revised version: 2.3.1 Acceptance bias

Earlier version: Motivational tendency was best predicted by a model with one component, with its highest factor loadings from psychiatric measures (increasing values indicate symptom severity; SHAPS71: -0.665; AES70: -0.588; Dimensional Anhedonia Rating Scale [DARS]73: -0.487). Weaker loadings were found for circadian measures (higher values indicate later chronotype; Mornigness-Eveningness Questionnaire [MEQ]11: - 0.262; Munich Chronotype Questionnaire [MCTQ]74: -0.117) and metabolic measures (higher values indicate higher metabolic risk; body mass index [BMI]: -0.115; Finnish Type-2 Diabetes Risk Score questionnaire [FINDRISC]75: -0.068). Permutation testing indicated the predictive value of the resulting component (with factor loadings described above) was significant out-of-sample (root-mean-squared error [RMSE]=0.203, p=.001).

Revised version: The acceptance bias was best predicted by a model with one component, with its highest factor loadings from psychiatric measures (increasing values indicate symptom severity; SHAPS71: -0.665; AES70: -0.588; Dimensional Anhedonia Rating Scale [DARS]73: -0.487). Weaker loadings were found for circadian measures (higher values indicate later chronotype; Mornigness-Eveningness Questionnaire [MEQ]11: - 0.262; Munich Chronotype Questionnaire [MCTQ]74: -0.117) and metabolic measures (higher values indicate higher metabolic risk; body mass index [BMI]: -0.115; Finnish Type-2 Diabetes Risk Score questionnaire [FINDRISC]75: -0.068). Permutation testing indicated the predictive value of the resulting component (with factor loadings described above) was significant out-of-sample (root-mean-squared error [RMSE]=0.203, p=.001).

Earlier version: Bayesian GLMs confirmed evidence for psychiatric questionnaire measures predicting motivational tendency (SHAPS: M=-0.109; 95% highest density interval (HDI)=[-0.17,-0.04]; AES: M=-0.096; 95%HDI=[-0.15,-0.03]; DARS: M=-0.061; 95%HDI=[-0.13,-0.01]; Fig. 4A). Post-hoc GLMs on DARS sub-scales showed an effect for the sensory subscale (M=-0.050; 95%HDI=[-0.10,-0.01]). This result of neuropsychiatric symptoms predicting a lower motivational tendency is in line with our pre-registered hypothesis. For the MEQ (95%HDI=[-0.09,0.06]), MCTQ (95%HDI=[-0.17,0.05]), BMI (95%HDI=[-0.19,0.01]), and FINDRISC (95%HDI=[- 0.09,0.03]) no relationship with motivational tendency was found, consistent with the smaller magnitude of reported component loadings from the PLS regression. This null finding for dimensional measures of circadian rhythm and metabolic health was not in line with our pre-registered hypotheses.

Revised version: Bayesian GLMs confirmed evidence for psychiatric questionnaire measures predicting acceptance bias (SHAPS: M=-0.109; 95% highest density interval (HDI)=[-0.17,-0.04]; AES: M=-0.096; 95%HDI=[-0.15,-0.03]; DARS: M=-0.061; 95%HDI=[-0.13,- 0.01]; Fig. 4A). Post-hoc GLMs on DARS sub-scales showed an effect for the sensory subscale (M=-0.050; 95%HDI=[-0.10,-0.01]). This result of neuropsychiatric symptoms predicting a lower acceptance bias is in line with our pre-registered hypothesis. For the MEQ (95%HDI=[-0.09,0.06]), MCTQ (95%HDI=[-0.17,0.05]), BMI (95%HDI=[-0.19,0.01]), and FINDRISC (95%HDI=[-0.09,0.03]) no relationship with acceptance bias was found, consistent with the smaller magnitude of reported component loadings from the PLS regression. This null finding for dimensional measures of circadian rhythm and metabolic health was not in line with our pre-registered hypotheses.

2.3.2 Effort sensitivity

For effort sensitivity, the intercept-only model outperformed models incorporating questionnaire predictors based on RMSE.

2.3.3 Reward sensitivity

For reward sensitivity, the intercept-only model outperformed models incorporating questionnaire predictors based on RMSE. This result was not in line with our pre-registered expectations.

2.3.4 Questionnaire measures predict model agnostic task measures

Both SHAPS (M=-0.07; 95%HDI=[-0.12,-0.03]) and AES (M=-0.05; 95%HDI=[-0.10,-0.002]) sum scores could predict the proportion of accepted trials averaged across effort and reward levels (Fig. S4).

Associations between task parameter estimates and psychiatric measures.

Earlier version: A: Visualizations of associations between the motivational tendency task parameter and the Snaith-Hamilton Pleasure Scale (SHAPS), the Dimensional Anhedonia Rating Scale (DARS)73, and the Apathy Evaluation Scale (AES)70 B-C: Comparison of motivational tendency (left) and effort sensitivity (right) between a sample of participants meeting criteria for current major depressive disorder (MDD; purple, upper) on the the Mini-International Neuropsychiatric Interview 7.0.1 (M.I.N.I)76 and age- and gender-matched controls (yellow, lower).

Revised version: A: Visualizations of associations between the acceptance bias task parameter and the Snaith-Hamilton Pleasure Scale (SHAPS), the Dimensional Anhedonia Rating Scale (DARS)73, and the Apathy Evaluation Scale (AES)70 B-C: Comparison of acceptance bias (left) and effort sensitivity (right) between a sample of participants meeting criteria for current major depressive disorder (MDD; purple, upper) on the the Mini-International Neuropsychiatric Interview 7.0.1 (M.I.N.I)76 and age- and gender-matched controls (yellow, lower).

2.4 Diagnostic analysis: Depressed and healthy subjects differ in effort-based decision-making

Earlier version: In an exploratory analysis, we compared a sample of N=56 participants that met criteria for current major depressive disorder (MDD), to fifty-six healthy controls (HC), matched by age (MDD: M=37.07; HC: M=37.09, p=.99) and gender (MDD: 31 female, 23 male, 2 non-binary; HC: 32 female, 22 male, 2 non-binary; p=.98). Effort-discounting effects were confirmed in both groups. For both groups, model fiding and comparison identified the full parabolic model as the best-fiding model. We used age- and gender-controlled Bayesian GLMs to compare individual-level mean posterior parameter values between groups.

Revised version: In an exploratory analysis, we compared a sample of N=56 participants that met criteria for current major depressive disorder (MDD), to fifty-six healthy controls (HC), matched by age (MDD: M=37.07; HC: M=37.09, p=.99) and gender (MDD: 31 female, 23 male, 2 non-binary; HC: 32 female, 22 male, 2 non-binary; p=.98). Effort-discounting effects were confirmed in both groups. For both groups, model fitting and comparison identified the full parabolic model as the best-fitting model. We used age- and gender-controlled Bayesian GLMs to compare individual-level mean posterior parameter values between groups.

Earlier version: 2.4.1 Motivational tendency

Revised version: 2.4.1 Acceptance bias

Earlier version: As in our transdiagnostic analyses of continuous neuropsychiatric measures (Results 2.3), we found evidence for a lower motivational tendency parameter in the MDD group compared to HCs (M=-0.111, 95% HDI=[ -0.20,-0.03]) (Fig. 4B). This result confirmed our pre-registered hypothesis.

Revised version: As in our transdiagnostic analyses of continuous neuropsychiatric measures (Results 2.3), we found evidence for a lower acceptance bias parameter in the MDD group compared to HCs (M=-0.111, 95% HDI=[-0.20,-0.03]) (Fig. 4B). This result confirmed our pre-registered hypothesis.

2.4.2 Effort sensitivity

Unlike our transdiagnostic analyses, we also found evidence for lower effort sensitivity in the MDD group compared to HCs (M=-0.111, 95% HDI=[-0.22,-0.01]) (Fig. 4C).

2.4.3 Reward sensitivity

There was no evidence for a group difference in reward sensitivity (95%HDI=[-0.07,0.11]), as in our transdiagnostic analyses.

2.5 Circadian measures affect effort-based decision-making

Due to our hypothesised interaction between circadian preference and time-of-day, testing was conducted in two specified time windows: morning (08:00–11:59) and evening (18:00–21:59), resulting in a binary time-of-day measure (morning vs. evening testing). A total of 492 participants completed the study in the morning testing window and 458 in the evening testing window. We used the two chronotype questionnaires to identify two established circadian phenotypes: “early” or “late” chronotype (see Methods 4.5), behavioural categories indicating underlying chronobiological differences9,11,74. These classifications result in four sub-sample groups, with 89 early chronotypes (morning testing: n=63; evening testing: n=26) and 75 late chronotypes (morning testing: n=20; evening testing: n=55).

Earlier version: Bayesian GLMs, controlling for age and gender, predicting task parameters by time-of-day and chronotype showed effects of chronotype on reward sensitivity (i.e. those with a late chronotype had a higher reward sensitivity; M= 0.325, 95% HDI=[0.19,0.46]) and motivational tendency (higher motivational tendency in early chronotypes; M=-0.248, 95% HDI=[-0.37,-0.11]), as well as an interaction between chronotype and time-of-day on motivational tendency (M=0.309, 95% HDI=[0.15,0.48]).

Revised version: Bayesian GLMs, controlling for age and gender, predicting task parameters by time-of-day and chronotype showed effects of chronotype on reward sensitivity (i.e. those with a late chronotype had a higher reward sensitivity; M=0.325, 95% HDI=[0.19,0.46]) and acceptance bias (higher acceptance bias in early chronotypes; M=-0.248, 95% HDI=[-0.37,-0.11]), as well as an interaction between chronotype and time-of-day on acceptance bias (M=0.309, 95% HDI=[0.15,0.48]).

2.5.1 Additional pre-registered data collection

Earlier version: As these analyses rely on unevenly distributed sub-samples, we conducted an additional, pre-registered data collection to replicate and extend these findings (hdps://osf.io/y4ke). We screened participants for their chronotype and then invited early chronotypes to take part in our study in the evening testing window, and late chronotypes in the morning testing window (Methods 4.5.1).

Revised version: As these analyses rely on unevenly distributed sub-samples, we conducted an additional, pre-registered data collection to replicate and extend these findings (https://osf.io/y4fbe). We screened participants for their chronotype and then invited early chronotypes to take part in our study in the evening testing window, and late chronotypes in the morning testing window (Methods 4.5.1).

Using our pre-registered Bayesian stopping rule, we tested 13 early chronotype participants and 20 late chronotype participants. The data was then combined with the data from our main data collection, resulting in a full sample of n=197 participants that was used for subsequent chronotype analyses (see Table 2 for sample characteristics and statistical significance of differences).

Added in the revised version: 2.5.2 Acceptance bias

Removed in the revised version: 2.5.2 Motivational tendency

Earlier version: Late chronotypes showed a lower motivational tendency than early chronotypes (M=-0.11, 95% HDI=[-0.22,-0.02])—comparable to effects of transdiagnostic measures of apathy and anhedonia, as well as diagnostic criteria for depression. Crucially, we found motivational tendency was modulated by an interaction between chronotype and time-of-day (M=0.19, 95% HDI=[0.05,0.33]): post-hoc GLMs in each chronotype group showed this was driven by a time-of-day effect within late, rather than early, chronotype participants (M=0.12, 95% HDI=[0.02,0.22], such that late chronotype participants showed a lower motivational tendency in the morning testing sessions, and a higher motivational tendency in the evening testing sessions; early chronotype: 95% HDI=[-0.16,0.04]) (Fig. 5A). These results of a main effect and an interaction effect of chronotype on motivational tendency confirmed our pre-registered hypothesis.

Revised version: Late chronotypes showed a lower acceptance bias than early chronotypes (M=-0.11, 95% HDI=[-0.22,-0.02])—comparable to effects of transdiagnostic measures of apathy and anhedonia, as well as diagnostic criteria for depression. Crucially, we found acceptance bias was modulated by an interaction between chronotype and time-of-day (M=0.19, 95% HDI=[0.05,0.33]): post-hoc GLMs in each chronotype group showed this was driven by a time-of-day effect within late, rather than early, chronotype participants (M=0.12, 95% HDI=[0.02,0.22], such that late chronotype participants showed a lower acceptance bias in the morning testing sessions, and a higher acceptance bias in the evening testing sessions; early chronotype: 95% HDI=[-0.16,0.04]) (Fig. 5A). These results of a main effect and an interaction effect of chronotype on acceptance bias confirmed our pre-registered hypothesis.

Earlier version: 2.5.2.1 Neuropsychiatric symptoms and circadian measures have separable effects on motivational tendency

Revised version: 2.5.2.1 Neuropsychiatric symptoms and circadian measures have separable effects on acceptance bias

Earlier version: Exploratory analyses testing for the effects of neuropsychiatric questionnaires on motivational tendency in the subsamples of early and late chronotypes confirmed the predictive value of the SHAPS (M=-0.24, 95% HDI=[-0.42,-0.06]), the DARS (M=- 0.16, 95% HDI=[-0.31,-0.01]), and the AES (M=-0.18, 95% HDI=[-0.32,-0.02]) on motivational tendency.

Revised version: Exploratory analyses testing for the effects of neuropsychiatric questionnaires on acceptance bias in the subsamples of early and late chronotypes confirmed the predictive value of the SHAPS (M=-0.24, 95% HDI=[-0.42,-0.06]), the DARS (M=-0.16, 95% HDI=[-0.31,-0.01]), and the AES (M=-0.18, 95% HDI=[-0.32,-0.02]) on acceptance bias.

Earlier version: For the SHAPS, we find that when adding the measures of chronotype and time-of-day back into the GLMs, the main effect of the SHAPS (M=-0.26, 95% HDI=[- 0.43,-0.07]), the main effect of chronotype (M=-0.11, 95% HDI=[-0.22,-0.01]), and the interaction effect of chronotype and time-of-day (M=0.20, 95% HDI=[0.07,0.34]) on motivational tendency remain. Model comparison by LOOIC reveals motivational tendency is best predicted by the model including the SHAPS, chronotype and time- of-day as predictors, followed by the model including only the SHAPS. Note that this approach to model comparison penalizes models for increasing complexity.

Revised version: For the SHAPS, we find that when adding the measures of chronotype and time-of-day back into the GLMs, the main effect of the SHAPS (M=-0.26, 95% HDI=[-0.43,-0.07]), the main effect of chronotype (M=-0.11, 95% HDI=[-0.22,-0.01]), and the interaction effect of chronotype and time-of-day (M=0.20, 95% HDI=[0.07,0.34]) on acceptance bias remain. Model comparison by LOOIC reveals acceptance bias is best predicted by the model including the SHAPS, chronotype and time-of-day as predictors, followed by the model including only the SHAPS. Note that this approach to model comparison penalizes models for increasing complexity.

Earlier version: Repeating these steps with the DARS, the main effect of the DARS is found numerically, but the 95% HDI just includes 0 (M=-0.15, 95% HDI=[-0.30,0.002]). The main effect of chronotype (M=-0.11, 95% HDI=[-0.21,-0.01]), and the interaction effect of chronotype and time-of-day (M=0.18, 95% HDI=[0.05,0.33]) on motivational tendency remain. Model comparison identifies the model including the DARS and circadian measures as the best model, followed by the model including only the DARS.

Revised version: Repeating these steps with the DARS, the main effect of the DARS is found numerically, but the 95% HDI just includes 0 (M=-0.15, 95% HDI=[-0.30,0.002]). The main effect of chronotype (M=-0.11, 95% HDI=[-0.21,-0.01]), and the interaction effect of chronotype and time-of-day (M=0.18, 95% HDI=[0.05,0.33]) on acceptance bias remain. Model comparison identifies the model including the DARS and circadian measures as the best model, followed by the model including only the DARS.

Earlier version: For the AES, the main effect of the AES is found (M=-0.19, 95% HDI=[-0.35,- 0.04]). For the main effect of chronotype, the 95% narrowly includes 0 (M=-0.10, 95% HDI=[-0.21,0.002]), while the interaction effect of chronotype and time-of-day (M=0.20, 95% HDI=[0.07,0.34]) on motivational tendency remains. Model comparison identifies the model including the AES and circadian measures as the best model, followed by the model including only the AES.

Revised version: For the AES, the main effect of the AES is found (M=-0.19, 95% HDI=[-0.35,- 0.04]). For the main effect of chronotype, the 95% narrowly includes 0 (M=-0.10, 95% HDI=[-0.21,0.002]), while the interaction effect of chronotype and time-of-day (M=0.20, 95% HDI=[0.07,0.34]) on acceptance bias remains. Model comparison identifies the model including the AES and circadian measures as the best model, followed by the model including only the AES.

2.5.3 Effort sensitivity

We found no evidence for circadian or time-of-day effects on effort sensitivity (chronotype main effect: 95%HDI=[-0.06,0.18], time-of-day main effect: 95%HDI=[- 0.08,0.13]).

2.5.4 Reward sensitivity

Participants with an early chronotype had a lower reward sensitivity parameter than those with a late chronotype (M=0.27, 95% HDI=[0.16,0.38]). We found no effect of time-of-day on reward sensitivity (95%HDI=[-0.09,0.11]) (Fig. 5B). These results were in line with our pre-registered hypotheses.

Effects of chronotype and time-of-day on task parameter estimates.

Earlier version: A: Effect of chronotype and time-of-day on reward sensitivity parameter estimates. B: Effect of chronotype and time-of-day on motivational tendency parameter estimates.

Revised version: A: Effect of chronotype and time-of-day on reward sensitivity parameter estimates. B: Effect of chronotype and time-of-day on acceptance bias parameter estimates.

3 Discussion

Earlier version: Various neuropsychiatric disorders are marked by disruptions in circadian rhythm, such as a late chronotype. However, research has rarely investigated how transdiagnostic mechanisms underlying neuropsychiatric conditions may relate to inter-individual differences in circadian rhythm. Here, combining a large-scale online study with computational modelling, we replicate and extend previous work linking anhedonia, apathy, and depression to a lower motivational tendency to exert effort for reward. Crucially, we found participants with a late compared to early chronotype show the same decrease in motivational tendency. Moreover, by testing participants at chronotype-compatible and -incompatible times of day, we discovered late chronotypes show a decreased motivational tendency to exert effort for reward when tested in the morning compared to evening. This reveals neuropsychiatric symptoms and chronotype (interacting with time-of-testing) show paralleling effects on effort-based decision-making. Our results demonstrate a crucial role for circadian rhythm in computational psychiatry, potentially affecting our assessment and treatment of neurocognitive mechanisms.

Revised version: Various neuropsychiatric disorders are marked by disruptions in circadian rhythm, such as a late chronotype. However, research has rarely investigated how transdiagnostic mechanisms underlying neuropsychiatric conditions may relate to inter-individual differences in circadian rhythm. Here, combining a large-scale online study with computational modelling, we replicate and extend previous work linking anhedonia, apathy, and depression to a lower bias to accept effort for reward. Crucially, we found participants with a late compared to early chronotype show the same decrease in acceptance bias. Moreover, by testing participants at chronotypecompatible and – incompatible times of day, we discovered late chronotypes show a decreased acceptance bias to exert effort for reward when tested in the morning compared to evening. This reveals neuropsychiatric symptoms and chronotype (interacting with time-of-testing) show paralleling effects on effort-based decision-making. Our results demonstrate a crucial role for circadian rhythm in computational psychiatry, potentially affecting our assessment and treatment of neurocognitive mechanisms.

Earlier version: We replicate and extend effects of aberrant effort-based decision-making in neuropsychiatric syndromes in a large, broadly population-representative sample. Our finding that dimensional measures of apathy and anhedonia predict motivational tendency, a computational parameter describing someone’s tendency to exert effort for reward, aligns with previous reports of impaired effort-based decision-making in psychiatric31,77–85 and neurodegenerative populations86–90, as well as studies linking effort-based decision-making with apathy and anhedonia specifically. The positive link between effort-based decision-making and apathy and anhedonia has been observed in both patients77,84,91 and healthy controls77,92,93, though some did not find this effect85.

Revised version: We replicate and extend effects of aberrant effort-based decision-making in neuropsychiatric syndromes in a large, broadly population-representative sample. Our finding that dimensional measures of apathy and anhedonia predict acceptance bias, a computational parameter describing someone’s tendency to exert effort for reward, aligns with previous reports of impaired effort-based decision-making in psychiatric31,77–85 and neurodegenerative populations86–90, as well as studies linking effort-based decision-making with apathy and anhedonia specifically. The positive link between effort-based decision-making and apathy and anhedonia has been observed in both patients77,84,91 and healthy controls77,92,93, though some did not find this effect85.

Earlier version: Our work supports previous theories that impaired effort-based decision-making represents a common, transdiagnostic mechanism across the psychiatric and neurological syndromes of anhedonia and apathy (respectively). We found corresponding effects of apathy and anhedonia on the same computational parameter — motivational tendency — reinforcing the suggestion of possible shared mechanistic underpinnings of the two motivational syndromes24. Aberrant effort-based decision-making may manifest behaviourally as deficient motivation, a symptom category that cuts across traditional disease boundaries of psychiatric, neurological, and neurodevelopmental disorders32.

Revised version: Our work supports previous theories that impaired effort-based decision-making represents a common, transdiagnostic mechanism across the psychiatric and neurological syndromes of anhedonia and apathy (respectively). We found corresponding effects of apathy and anhedonia on the same computational parameter - acceptance bias – reinforcing the suggestion of possible shared mechanistic underpinnings of the two motivational syndromes24. Aberrant effort-based decision-making may manifest behaviourally as deficient motivation, a symptom category that cuts across traditional disease boundaries of psychiatric, neurological, and neurodevelopmental disorders32.

Earlier version: Our categorical (diagnostic-criteria based) analysis comparing depressed to healthy subjects likewise found depressed patients showed a lower motivational tendency, echoing our dimensional results in apathy and anhedonia. In addition, our categorical analysis revealed a distinct effect of group on effort sensitivity: depressed subjects had lower effort sensitivity, meaning their decisions were less influenced by effort changes. Possibly, this effect stems from decreased perceived differences in effort levels, as recently reported94, indicating there are both dimensional (transdiagnostic) and potentially some diagnosis-specific effects of mental health on effort-based decision-making.

Revised version: Our categorical (diagnostic-criteria based) analysis comparing depressed to healthy subjects likewise found depressed patients showed a lower acceptance bias, echoing our dimensional results in apathy and anhedonia. In addition, our categorical analysis revealed a distinct effect of group on effort sensitivity: depressed subjects had lower effort sensitivity, meaning their decisions were less influenced by effort changes. Possibly, this effect stems from decreased perceived differences in effort levels, as recently reported94, indicating there are both dimensional (transdiagnostic) and potentially some diagnosis-specific effects of mental health on effort-based decision-making.

Earlier version: It is possible that a higher motivational tendency reflects a more optimistic assessment of future task success, in line with work on the optimism bias95; however our task intentionally minimized unsuccessful trials by titrating effort and reward; future studies should explore this more directly.

Revised version: It is possible that a higher acceptance bias reflects a more optimistic assessment of future task success, in line with work on the optimism bias95; however our task intentionally minimized unsuccessful trials by titrating effort and reward; future studies should explore this more directly.

Earlier version: We also found circadian effects on effort-based decision-making, paralleling those of apathy, anhedonia, and depression measures in both the affected neurocomputational parameter and the direction of effect. We observed a difference in motivational tendency between chronotypes, with late chronotypes showing a lower tendency to accept to exert effort for reward. Previous studies have suggested late chronotypes were also less accepting of delays62 and risk for reward64,65.

Revised version: We also found circadian effects on effort-based decision-making, paralleling those of apathy, anhedonia, and depression measures in both the affected neurocomputational parameter and the direction of effect. We observed a difference in acceptance bias between chronotypes, with late chronotypes showing a lower tendency to accept to exert effort for reward. Previous studies have suggested late chronotypes were also less accepting of delays62 and risk for reward64,65.

Earlier version: Most importantly, we found an interaction between chronotype and time-of-day in a synchrony effect manner: early and late chronotypes showed a higher motivational tendency towards accepting effort for reward at their preferred time of day. This effect was driven by the late chronotype group, who showed a markedly lower motivational tendency in the morning, but a higher tendency in the evening. This suggests that chronotype effects on neurocomputational parameters such as motivational tendency depend on time-of-testing. Synchrony effects have previously been observed in other cognitive domains including inhibitory control96, adention97, learning98, and memory99. One interpretation of our cognitive synchrony effects may be that late chronotype participants show a diminished ability to adapt to suboptimal times-of-day due to reduced cognitive resources100.

Revised version: Most importantly, we found an interaction between chronotype and time-of-day in a synchrony effect manner: early and late chronotypes showed a higher acceptance bias towards accepting effort for reward at their preferred time of day. This effect was driven by the late chronotype group, who showed a markedly lower acceptance bias in the morning, but a higher tendency in the evening. This suggests that chronotype effects on neurocomputational parameters such as acceptance bias depend on time-of-testing. Synchrony effects have previously been observed in other cognitive domains including inhibitory control96, attention97, learning98, and memory99. One interpretation of our cognitive synchrony effects may be that late chronotype participants show a diminished ability to adapt to suboptimal times-of-day due to reduced cognitive resources100.

We also report a distinct effect of chronotype on effort-based decision-making that is not paralleled by effects of neuropsychiatric symptoms, nor dependent on time-of-day. Compared to early chronotypes, late chronotypes were more guided by differences in reward value, indicated by higher reward sensitivity parameters. Previous studies report altered reward functioning in late chronotypes, who show a reduced reactivity to reward in the medial prefrontal cortex, a key component of reward circuitry101–103. Note this is not incompatible with higher reward sensitivity due to our modelling approach, in which higher reward sensitivity does not imply higher reward valuation, but rather larger subjective value differences between reward levels. Therefore, reduced reactivity to reward could be compatible with late chronotypes devaluing low reward levels more, which in our models would emerge as a reduced reward sensitivity parameter.

Earlier version: It is striking that the effects of neuropsychiatric symptoms on effort-based decision-making largely are paralleled by circadian effects on the same neurocomputational parameter. Exploratory analyses predicting motivational tendency by neuropsychiatric symptoms and circadian measures simultaneously indicate the effects go beyond recapitulating each other, but rather explain separable parts of the variance in motivational tendency. Overall, our results raise the possibility of altered effort-reward processing as a critical mechanism linking neuropsychiatric conditions and circadian rhythm. Previous research demonstrated depressed patients with an evening chronotype show increased diurnal mood variation104. Our finding of time-of-day differences in motivational tendency among late chronotypes illustrates a potential cognitive underpinning for the observed diurnal characteristic within depression and late chronotype. Together, these findings support the idea of a circadian psychiatric phenotype14,17, which should be considered in measurement (e.g., design of computational psychiatry studies) and potentially treatment (e.g., administration of motivation-based psychological interventions, which could be timed compatibly with chronotype).

Revised version: It is striking that the effects of neuropsychiatric symptoms on effort-based decision-making largely are paralleled by circadian effects on the same neurocomputational parameter. Exploratory analyses predicting acceptance bias by neuropsychiatric symptoms and circadian measures simultaneously indicate the effects go beyond recapitulating each other, but rather explain separable parts of the variance in acceptance bias. Overall, our results raise the possibility of altered effort-reward processing as a critical mechanism linking neuropsychiatric conditions and circadian rhythm. Previous research demonstrated depressed patients with an evening chronotype show increased diurnal mood variation104. Our finding of time-of-day differences in acceptance bias among late chronotypes illustrates a potential cognitive underpinning for the observed diurnal characteristic within depression and late chronotype. Together, these findings support the idea of a circadian psychiatric phenotype14,17, which should be considered in measurement (e.g., design of computational psychiatry studies) and potentially treatment (e.g., administration of motivation-based psychological interventions, which could be timed compatibly with chronotype).

Earlier version: To our surprise, we did not find statistical evidence for a relationship between effort-based decision-making and measures of metabolic health (BMI and risk for type-2 diabetes). Our analyses linking BMI to motivational tendency reveal a numeric effect in line with our hypothesis: a higher BMI relating to a lower motivational tendency. However, the 95% HDI for this effect narrowly included zero (95%HDI=[-0.19,0.01]). Possibly, our sample did not have sufficient variance in metabolic health to detect dimensional metabolic effects in a current general population sample. A recent study by our group investigates the same neurocomputational parameters of effort-based decision-making in participants with type-2 diabetes and non-diabetic controls matched by age, gender, and physical activity105. We report a group effect on the motivational tendency parameter, with type-2 diabetic patients showing a lower tendency to exert effort for reward.

Revised version: To our surprise, we did not find statistical evidence for a relationship between effort-based decision-making and measures of metabolic health (BMI and risk for type-2 diabetes). Our analyses linking BMI to acceptance bias reveal a numeric effect in line with our hypothesis: a higher BMI relating to a lower acceptance bias. However, the 95% HDI for this effect narrowly included zero (95%HDI=[-0.19,0.01]). Possibly, our sample did not have sufficient variance in metabolic health to detect dimensional metabolic effects in a current general population sample. A recent study by our group investigates the same neurocomputational parameters of effort-based decision-making in participants with type-2 diabetes and non-diabetic controls matched by age, gender, and physical activity105. We report a group effect on the acceptance bias parameter, with type-2 diabetic patients showing a lower tendency to exert effort for reward.

Earlier version: Our study results should be considered in light of a few limitations. First, we used online self-report measures of neuropsychiatric symptoms and depression status. There has been a large shift toward online data collection in psychiatric research, and while online data is undoubtedly noisier, results (including our own, presented in the supplemental material) usually show excellent accordance with labbased studies106. Similarly, we lack biological measures of circadian rhythm, the gold standard of chronotype assessment. However, this concern might be mitigated by previous reports of high covariance between biological- and questionnaire-based circadian measures107,108, as well as significant chronobiological differences between the questionnaire-determined chronotypes10,109 we use in our key findings. Nevertheless, future work should incorporate biological measures in adempts to replicate circadian effects on effort-based decision-making. This could take the form of identifying chronotypes by DNA analysis or dim-light melatonin onset, or continuous measurements of circadian proxies, such as core body temperature, heart rate, or actigraphy.

Revised version: Our study results should be considered in light of a few limitations. First, we used online self-report measures of neuropsychiatric symptoms and depression status. There has been a large shift toward online data collection in psychiatric research, and while online data is undoubtedly noisier, results (including our own, presented in the supplemental material) usually show excellent accordance with lab-based studies106. Similarly, we lack biological measures of circadian rhythm, the gold standard of chronotype assessment. However, this concern might be mitigated by previous reports of high covariance between biological- and questionnaire-based circadian measures107,108, as well as significant chronobiological differences between the questionnaire-determined chronotypes10,109 we use in our key findings. Nevertheless, future work should incorporate biological measures in attempts to replicate circadian effects on effort-based decision-making. This could take the form of identifying chronotypes by DNA analysis or dim-light melatonin onset, or continuous measurements of circadian proxies, such as core body temperature, heart rate, or actigraphy.

Note also that our time-of-day effects are limited by a between-subjects study design (i.e., the same participants were not tested in morning and evening sessions). It will be interesting to explore such diurnal variation in effort-based decision-making within individuals. The newly developed effort-expenditure task we present here may lend itself particularly well to such endeavours. First, it allows remote testing, meaning subjects can complete the task at different times of the day without in-person testing. Second, we demonstrated good test-retest reliability of task measures when time of testing was held constant within participants. This good test-retest reliability of our task contrasts with recent reports of poor test-retest reliability of other tasks and computational modelling parameters110.

Our reported analyses investigating neuropsychiatric and circadian effects on effort-based decision-making simultaneously are exploratory, as our study design was not ideally set out to examine this. Further work is needed to disentangle separable effects of neuropsychiatric and circadian measures on effort-based decision-making. One approach could be a group-based study design enabling the dissociation of the two effects (e.g., examining high-anhedonia participants with early chronotypes and low-anhedonia patients with late chronotypes, as well as the respective other, more common groupings, and testing each group in the morning and evening to examine time-of-day interactions with both anhedonia and chronotype).

Finally, we would like to note that as our study is based on a general population sample, rather than a clinical one. Hence, we cannot speak to transdiagnosticity on the level of multiple diagnostic categories.

Taken together, our results implicate circadian rhythm as an important factor in effort-based decision-making and its relationship to neuropsychiatric conditions. These results have implications for research, clinical interventions, and policy. We demonstrate that neuropsychiatric effects on effort-based decision-making are paralleled by effects of circadian rhythm and time-of-day. Exploratory analyses suggest these effects account for separable parts of the variance in effort-based decision-making. It unlikely that effects of neuropsychiatric effects on effort-based decision-making reported here and in previous literature are a spurious result due to multicollinearity with chronotype. Yet, not accounting for chronotype and time of testing, which is the predominant practice in the field, could affect results. This could take the form of either inflating or depressing results in the existing literature. On the one hand, reported neuropsychiatric effects may be inflated by systematic circadian differences between participants (i.e., overrepresentation of late chronotype in patient samples), which could be further amplified by time of testing (often the morning, incompatible with late chronotypes, and producing motivational impairments on neurocognitive measures). On the other hand, true effects may be masked by interactions between chronotype and time-of-day: Testing psychiatric subjects with a late chronotype in the evening (e.g., as a consequence of subject-selected testing times) may paint a false picture of group equivalence, as researchers are only observing part of a daily trajectory.

Our growing understanding of the relationship between circadian rhythm and neuropsychiatry may allow for critical advances in improving therapeutic outcomes from treatments111,112. Such advances are particularly called for in the case of symptoms of apathy and anhedonia, as current treatments often fail to improve motivational deficits27–29, but could potentially be coupled with a patient’s chronotype to increase efficacy. At minimum, clinical trials predicting change in motivational measures, such as effort-based decision-making, should assess patients at similar times of day, as this could reduce or inflate treatment effects.

Circadian rhythm and neuropsychiatric syndromes may affect motivation via parallel, as well as distinct, mechanisms—but crucially, this overlap is dependent on time of testing. Our work suggests that chronotype and time of testing are essential variables to consider in future effort-based decision-making experiments, particularly those measuring effort-based decision-making in patient groups, such as those with depression, high apathy, or high anhedonia. Beyond experimental work, future interventions should consider the role of chronotype in measurement and modulation of motivation.

4 Methods and materials

4.1 Study protocol

Earlier version: After providing demographics and basic medical history, subjects completed an effort-expenditure task, followed by a badery of self-report questionnaires. The study was coded in JavaScript, using Phaser v.3.50.0 for the task and jsPsych113 for questionnaires.

Revised version: After providing demographics and basic medical history, subjects completed an effort-expenditure task, followed by a battery of self-report questionnaires. The study was coded in JavaScript, using Phaser v.3.50.0 for the task and jsPsych113 for questionnaires.

4.1.1 Ethics

Earlier version: This study was approved by the University of Cambridge Human Biology Research Ethics Commidee (HBREC.2020.40). Participants provided informed consent through an online form, complying with the University of Cambridge Human Biology Research Ethics Commidee procedures for online studies.

Revised version: This study was approved by the University of Cambridge Human Biology Research Ethics Committee (HBREC.2020.40). Participants provided informed consent through an online form, complying with the University of Cambridge Human Biology Research Ethics Committee procedures for online studies.

4.1.2 Recruitment

We recruited participants using Prolific114, in September 2022. Data was collected on weekdays, in specified daily time-windows (morning testing: 08:00–11:59; evening testing: 18:00- 21:59). To sample participants broadly representative of the UK population in age, sex, and history of psychiatric disorder, we implemented a previously-described procedure67 using Prolific pre-screeners to obtain batches of participants aimed to match target numbers calculated based on UK population data.

Nine hundred and ninety-four participants completed all components and were paid a fixed rate of £6. A bonus of £10 was paid to ten participants. Subjects were told they could increase their chances of winning the bonus by engaging well with the study (e.g., reading questions carefully, following task instructions).

4.1.3 Effort-expenditure task

Earlier version: We developed a new effort-expenditure task that allowed us to assess effort-based decision-making in a remote seding; this task was also tested in-person to assess test-retest reliability. To increase engagement, we gamified the task to take place in an underwater seding and each challenge is framed as a race in which an octopus catches a shrimp. The task structure is shown in Figure 2A and the trial-level structure in Figure 2B.

Revised version: We developed a new effort-expenditure task that allowed us to assess effort-based decision-making in a remote setting; this task was also tested in-person to assess test-retest reliability. To increase engagement, we gamified the task to take place in an underwater setting and each challenge is framed as a race in which an octopus catches a shrimp. The task structure is shown in Figure 2A and the trial-level structure in Figure 2B.

The task began with an individual calibration phase to standardise maximum effort capacity, followed by the main task, which used a semi-adaptive staircase design to maximise the informative value of each choice.

For the calibration, subjects were prompted to collect points by clicking as fast as possible for ten seconds, repeated three times. The second and third repetitions were then averaged to serve as the maximum clicking capacity reference for the main task. One calibration trial was repeated at the end of the main task to monitor any notable changes in clicking capacity. Then, subjects were familiarized with their individually calibrated effort levels during a practice phase of the task. Effort levels were scaled to a given participant’s mean clicking speed (based on the calibration phase), and the time clicking must be sustained for. We used four effort-levels, corresponding to a clicking speed at 30% of a participant’s maximal capacity for 8 seconds (level 1), 50% for 11 seconds (level 2), 70% for 14 seconds (level 3), and 90% for 17 seconds (level 4). Therefore, in each trial, participants had to fulfil a certain number of mouse clicks (dependent on their capacity and the effort level) in a specific time (dependent on the effort level). Subjects were instructed to make mouse-clicks with the finger they normally use, and to not change fingers throughout the task (compliance was checked at the end of the main task). In the practice phase, all effort-levels were completed without reward associations, and failed levels were repeated until subjects succeed at each level. If a subject failed a level twice, the clicking capacity reference was adjusted to the speed reached in the practice. Finally, subjects needed to pass a six-question quiz to ensure task instructions were fully understood. If a subject failed any question on the quiz, they were returned to the instruction screens and re-took the quiz until all questions were answered correctly.

The main task took a binary-choice design: In each trial, participants accepted or rejected a challenge associated with one of the four specific effort levels and rewards. Reward was conceptualized as points (shrimp caught by the octopus) that could be collected in that trial. The points to win per challenge varied between four levels (2,3,4, or 5 points). If a subject accepted a given challenge, they needed to achieve the given effort-level to win the associated points. If a subject rejected a given challenge, they waited and received one point, with waiting times matched to the respective effort level to prevent confounding with delay discounting. Participants were able to infer their clicking progress from the distance between the octopus and the shrimp and the remaining time was indicated by a time-bar.

Earlier version: Subjects completed 64 trials, split into four blocks of 16 trials. For each subject, trial-by-trial presentation of effort-reward combinations were made semi-adaptively by 16 randomly interleaved staircases. Each of the 16 possible offers (4 effort-levels x 4 reward-levels) served as the starting point of one of the 16 staircase. Within each staircase, after a subject accepted a challenge, the next trial’s offer on that staircase was adjusted (by increasing effort or decreasing reward). After a subject rejected a challenge, the next offer on that staircase was adjusted by decreasing effort or increasing reward. This ensured subjects received each effort-reward combination at least once (as each participant completed all 16 staircases), while individualizing trial presentation to maximize the trials’ informative value. Therefore, in practice, even in the case of a subject rejecing all offers (and hence the staircasing procedures always adapting by decreasing effort or increasing reward), the full range of effort-reward combinations will be represented in the task across the startingpoints of all staircases (and therefore before adaption takeplace).

Revised version: Subjects completed 64 trials, split into four blocks of 16 trials. For each subject, trial-by-trial presentation of effort-reward combinations were made semi-adaptively by 16 randomly interleaved staircases. Each of the 16 possible offers (4 effort-levels × 4 reward-levels) served as the starting point of one of the 16 staircase. Within each staircase, after a subject accepted a challenge, the next trial’s offer on that staircase was adjusted (by increasing effort or decreasing reward). After a subject rejected a challenge, the next offer on that staircase was adjusted by decreasing effort or increasing reward. This ensured subjects received each effort-reward combination at least once (as each participant completed all 16 staircases), while individualizing trial presentation to maximize the trials’ informative value. Therefore, in practice, even in the case of a subject rejecing all offers (and hence the staircasing procedures always adapting by decreasing effort or increasing reward), the full range of effort-reward combinations will be represented in the task across the startingpoints of all staircases (and therefore before adaption takeplace).

4.1.4 Self-report questionnaires

Earlier version: Subjects completed a questionnaire badery assessing mental and physical health, presented in a randomised order. We assessed anhedonia using the SHAPS71, as well as the DARS73. Apathy was assessed with the AES70. Additionally, we screened participants for meeting diagnostic criteria for current, past, or recurrent MDD using the M.I.N.I.76. Two questionnaires targeted circadian rhythm: the MEQ11 and the MCTQ74. Metabolic health was assessed by collecting self-reported height and weight, used to calculate BMI. Additionally, the FINDRISC75 was used to calculate individual risk scores for metabolic disease. Finally, the International Physical Activity Questionnaire (IPAQ)116 was included for exploratory investigations of physical activity.

Revised version: Subjects completed a questionnaire battery assessing mental and physical health, presented in a randomised order. We assessed anhedonia using the SHAPS71, as well as the DARS73. Apathy was assessed with the AES70. Additionally, we screened participants for meeting diagnostic criteria for current, past, or recurrent MDD using the M.I.N.I.76. Two questionnaires targeted circadian rhythm: the MEQ11 and the MCTQ74. Metabolic health was assessed by collecting self-reported height and weight, used to calculate BMI. Additionally, the FINDRISC75 was used to calculate individual risk scores for metabolic disease. Finally, the International Physical Activity Questionnaire (IPAQ)116 was included for exploratory investigations of physical activity.

4.1.5 Compliance checks and exclusion criteria

All exclusion criteria were preregistered. Participants were excluded when reporting a severe neurological condition (n=14) or English proficiency below B2 (i.e., good command/working knowledge; n=2).

To check compliance with the questionnaires, four catch questions were presented during questionnaires, including two easy questions (e.g., “Please answer ‘Not at all’.”) and two harder questions (e.g., “In the past week, I (would have) wanted to eat mouldy food.”, expected answer “Disagree” or “Definitely disagree”). Participants failing at least one easy question, or both harder questions were excluded (n=12).

Earlier version: As task-based exclusion criteria, subjects rejecting all offers were excluded (n=0). Participants had to have a clicking-calibration score of at least seven, as values below would lead to challenges with just one mouse-click (n=4). Subjects showing a large difference between minimum and maximum clicking speed (i.e., >3 standard deviations (SD)) during calibration trials were excluded, as a misestimation of the calibration reference is likely (n=3). Finally, subjects showing a large change in their clicking capacity (i.e., >3 SD) pre- to post-task were excluded, as it can be assumed the applied calibration was not valid during the task (n=1). We also asked two open-answer questions after completion of the main task to monitor participants’ self-reported task strategies as a way of assessing rule adherence.

Revised version: As task-based exclusion criteria, subjects rejecting all offers were excluded (n=0). Participants had to have a clicking-calibration score of at least seven, as values below would lead to challenges with just one mouse-click (n=4). Subjects showing a large difference between minimum and maximum clicking speed (i.e., >3 standard deviations (SD)) during calibration trials were excluded, as a misestimation of the calibration reference is likely (n=3). Finally, subjects showing a large change in their clicking capacity (i.e., >3 SD) pre- to post-task were excluded, as it can be assumed the applied calibration was not valid during the task (n=1). We also asked two openanswer questions after completion of the main task to monitor participants’ self-reported task strategies as a way of assessing rule adherence.

4.2 Analyses of effort-expenditure task data

4.2.1 Model-agnostic analyses

Using the proportion of accepted challenges as the dependent variable, we investigated main effects of effort- and reward-levels and their interaction, using a repeated measures analysis of variances (ANOVA) of repeated measures. This approach accommodates the unbalanced design resulting from the implemented staircasing procedure.

4.2.2 Model-based analyses

4.2.2.1 Model space

To model effort-based decision-making, we considered a model space of nine models. All models are variations of the economic decision-theory model, consisting of two basic equations. First, a cost function transforms costs and rewards associated with an action into a subjective value (SV): SV=(βR⋅R)−(βE⋅E) (1)

Earlier version: with βR and βE for reward and effort sensitivity, and ) and + for reward and effort. Higher effort and reward sensitivity mean the SV is more strongly influenced by changes in effort and reward, respectively (Fig. 3B-C). Hence, low effort and reward sensitivity mean the SV, and with that decision-making, is less guided by effort and reward offers, as would be in random decision-making.

Revised version: with βR and βE for reward and effort sensitivity, and R and E for reward and effort. Higher effort and reward sensitivity mean the SV is more strongly influenced by changes in effort and reward, respectively (Fig. 3B-C). Hence, low effort and reward sensitivity mean the SV, and with that decision-making, is less guided by effort and reward offers, as would be in random decision-making.

This SV is then transformed to an acceptance probability by a softmax function: p(accept)=11+e−(α+SV) (2)

Earlier version: with p(accept) for the predicted acceptance probability and 3 for the intercept representing motivational tendency. A high motivational tendency means a subjects has a tendency, or bias, to accept rather than reject offers (Fig. 3D).

Revised version: with p(accept) for the predicted acceptance probability and α for the intercept representing acceptance bias. A high acceptance bias means a subject has a bias, or tendency, to accept rather than reject effortful offers for reward (Fig. 3D).

Earlier version: The models differed in two aspects. First, inclusion or exclusion of the free parameters reward sensitivity (βR) and motivational tendency (α). Second, the form of the cost function, which used either a linear function (proportional discounting at all effort-levels), a parabolic function (increases at higher effort-levels are discounted over-proportionally) or an exponential function (increases at lower effort-lower levels are discounted over-proportionally). See supplement 1.1 for mathematical definitions of all models.

Revised version: The models differed in two aspects. First, inclusion or exclusion of the free parameters reward sensitivity (βR) and acceptance bias (α). Second, the form of the cost function, which used either a linear function (proportional discounting at all effort-levels), a parabolic function (increases at higher effort-levels are discounted over-proportionally) or an exponential function (increases at lower effort-lower levels are discounted over-proportionally). See supplement 1.1 for mathematical definitions of all models.

4.2.2.2 Model fitting, checks, and comparisons

Earlier version: We took a hierarchical Bayesian approach to model fiding117, implemented with the CmdStan R interface118, with Stan code adapted from hBayesDM119. Prior to model fiding, effort- and reward-levels were standardized for computational ease. All models were fit using Markov-Chain Monte Carlo (MCMC), with 2000 warm-up iterations and 6000 sampling iterations, by four chains. Model convergence and chain mixing were checked using numerical diagnostics of ESS and split R-hats, and by visually inspecting trace plots. We conducted parameter recoveries for all models, confirming their ability to meaningfully recover known parameters. Model performance was compared based on out-of-sample predictive accuracy using the LOOIC (lower is beder) and ELPD (higher is beder). The winning model was validated using posterior predictive checks, comparing model predictions to subjectwise observed choices.

Revised version: We took a hierarchical Bayesian approach to model fitting117, implemented with the CmdStan R interface118, with Stan code adapted from hBayesDM119. Prior to model fitting, effort- and reward-levels were standardized for computational ease. All models were fit using Markov-Chain Monte Carlo (MCMC), with 2000 warm-up iterations and 6000 sampling iterations, by four chains. Model convergence and chain mixing were checked using numerical diagnostics of ESS and split R-hats, and by visually inspecting trace plots. We conducted parameter recoveries for all models, confirming their ability to meaningfully recover known parameters. Model performance was compared based on out-of-sample predictive accuracy using the LOOIC (lower is better) and ELPD (higher is better). The winning model was validated using posterior predictive checks, comparing model predictions to subject wise observed choices.

4.2.2.3 Test-retest reliability

Earlier version: We conducted an in-person study to validate the effort-expenditure task and assess the test-retest reliability of our computational modelling parameters. A sample of N=30 participants was recruited and tested in two sessions, about one week apart. Test-retest reliability of task parameters was assessed by intra-class correlation coefficients, Pearson’s correlation coefficients (estimated both after model fiding and by embedding a correlation matrix into the model fiding procedure), and by testing the predictive accuracy of parameter estimates across sessions. See supplement 2 for full methods and results.

Revised version: We conducted an in-person study to validate the effort-expenditure task and assess the test-retest reliability of our computational modelling parameters. A sample of N=30 participants was recruited and tested in two sessions, about one week apart. Test-retest reliability of task parameters was assessed by intra-class correlation coefficients, Pearson’s correlation coefficients (estimated both after model fitting and by embedding a correlation matrix into the model fitting procedure), and by testing the predictive accuracy of parameter estimates across sessions. See supplement 2 for full methods and results.

4.3 Linking model parameters to outcome measures

To aid interpretability and comparability of effects, task parameters and questionnaire outcome measures were standardized to be between zero and one. Questionnaire measures resulting from the DARS, AES, and MEQ were additionally transformed to be interpretable with the same directionality within questionnaire groupings (i.e., for all psychiatric measures higher values are interpreted as higher symptom severity, for all circadian measures higher values are interpreted as later chronotype).

Earlier version: To investigate associations between effort-based decision-making and self-report questionnaires, we ran partial least squares (PLS) regressions with questionnaire outcome measures predicting modelling parameters. PLS regression allows joint modelling of questionnaire measures, without issues due to expected multicollinearity between questionnaires. Following the best practice of model validation120, data was split into a training (75%) and a testing (25%) subset. The training data was used to obtain the optimal number of components, based on ten-fold-cross validation, and to train the model. The winning model’s predictive performance was tested out-of-sample using the held-out testing data. Statistical significance of obtained effects (i.e., the predictive accuracy of the identified component and factor loadings) was assessed by permutation tests, probing the proportion of root-mean-squared errors (RMSEs) indicating stronger or equally strong predictive accuracy under the null hypothesis.

Revised version: To investigate associations between effort-based decision-making and self-report questionnaires, we ran partial least squares (PLS) regressions with questionnaire outcome measures predicting modelling parameters. PLS regression allows joint modelling of questionnaire measures, without issues due to expected multicollinearity between questionnaires. Following the best practice of model validation120, data was split into a training (75%) and a testing (25%) subset. The training data was used to obtain the optimal number of components, based on tenfold-cross validation, and to train the model. The winning model’s predictive performance was tested out-of-sample using the held-out testing data. Statistical significance of obtained effects (i.e., the predictive accuracy of the identified component and factor loadings) was assessed by permutation tests, probing the proportion of root-mean-squared errors (RMSEs) indicating stronger or equally strong predictive accuracy under the null hypothesis.

To follow up on relationships suggested by the PLS regression, we performed Bayesian general linear models (GLMs), adjusting for age and gender (male or female, imputing natal sex for non-binary participants, given low numbers).

4.4 Comparing depressed and healthy subjects

Earlier version: We compared participants meeting criteria for a current MDD based on the M.I.N.I.76, to a subset of age- and gender-matched healthy controls (HCs, participants that did not meet criteria for current MDD). For computational sparsity, we only fit the three best-fiding models from the full sample. Models were fit separately to the MDD and HC groups, using the same methods and parameters described above. Bayesian GLMs were used to quantify evidence for associations between individual-level modelling parameters and group status. As we could not be certain whether we would obtain a large enough sample size of subjects meeting criteria for MDD, these analyses were exploratory.

Revised version: We compared participants meeting criteria for a current MDD based on the M.I.N.I.76, to a subset of age- and gender-matched healthy controls (HCs, participants that did not meet criteria for current MDD). For computational sparsity, we only fit the three best-fitting models from the full sample. Models were fit separately to the MDD and HC groups, using the same methods and parameters described above. Bayesian GLMs were used to quantify evidence for associations between individual-level modelling parameters and group status. As we could not be certain whether we would obtain a large enough sample size of subjects meeting criteria for MDD, these analyses were exploratory.

4.5 Investigating circadian effects

We used the two circadian rhythm questionnaires to determine participants’ chronotypes. Early chronotype was defined as meeting criteria for “morning types” on the MEQ (MEQ sum score > 58)11 and having a midpoint of sleep on free days before 02:30. Late chronotype was defined as meeting criteria for “evening types” on the MEQ (MEQ sum score < 42) and having a midpoint of sleep on free days after 05:30. Subjects not falling into either category are categorized as intermediate chronotypes and were not included in these analyses.

We used Bayesian age- and gender-controlled GLMs to investigate effects of chronotype, time-of-day (morning- vs. evening-testing), and their interaction on subject-wise mean task parameters estimates.

4.5.1 Additional data collection

Earlier version: To improve the precision of estimated circadian effects on task parameters, we increased our sample size by conducting an additional pre-registered data collection (hdps://osf.io/y4ke). We implemented a screening study comprising the MEQ11 and MCTQ74. Taking the chronotyping approach described above, subjects with an early or late chronotype were identified. Early chronotypes were invited to take part in our study in the evening, late chronotypes in the morning.

Revised version: To improve the precision of estimated circadian effects on task parameters, we increased our sample size by conducting an additional pre-registered data collection (https://osf.io/y4fbe). We implemented a screening study comprising the MEQ11 and MCTQ74. Taking the chronotyping approach described above, subjects with an early or late chronotype were identified. Early chronotypes were invited to take part in our study in the evening, late chronotypes in the morning.

We implemented a Bayesian stopping rule to inform our data collection process, taking the following steps. First, participants were screened in batches of 250, and eligible participants were invited to the study session. Next, data resulting from this additional data collection was joined with data resulting from the main data collection and Bayesian GLMs were re-run, as described above. If our precision target of any 95% HDI reaching a maximum width 0.20 was met, we stopped data collection. Was the precision target not met, we returned to step one, and another batch of 250 participants was screened. In any case, data collection would be terminated once 200 eligible participants had completed the main study session.

Earlier version: 4.5.2 Differentiating between the effects of neuropsychiatric symptoms and circadian measures on motivational tendency

Revised version: 4.5.2 Differentiating between the effects of neuropsychiatric symptoms and circadian measures on acceptance bias

Earlier version: To investigate how the effects of neuropsychiatric symptoms on motivational tendency (2.3.1) relate to effects of chronotype and time-of-day on motivational tendency we conducted exploratory analyses. In the subsamples of participants with an early or late chronotype (including additionally collected data), we first ran Bayesian GLMs with neuropsychiatric questionnaire scores (SHAPS, DARS, AES respectively) predicting motivational tendency, controlling for age and gender. We next added an interaction term of chronotype and time-of-day into the GLMs, testing how this changes previously observed neuropsychiatric and circadian effects on motivational tendency. Finally, we conducted a model comparison using LOO, comparing between motivational tendency predicted by a neuropsychiatric questionnaire, motivational tendency predicted by chronotype and time-of-day, and motivational tendency predicted by a neuropsychiatric questionnaire and time-of-day (for each neuropsychiatric questionnaire, and controlling for age and gender).

Revised version: To investigate how the effects of neuropsychiatric symptoms on acceptance bias (2.3.1) relate to effects of chronotype and time-of-day on acceptance bias we conducted exploratory analyses. In the subsamples of participants with an early or late chronotype (including additionally collected data), we first ran Bayesian GLMs with neuropsychiatric questionnaire scores (SHAPS, DARS, AES respectively) predicting acceptance bias, controlling for age and gender. We next added an interaction term of chronotype and time-of-day into the GLMs, testing how this changes previously observed neuropsychiatric and circadian effects on acceptance bias. Finally, we conducted a model comparison using LOO, comparing between acceptance bias predicted by a neuropsychiatric questionnaire, acceptance bias predicted by chronotype and time-of-day, and acceptance bias predicted by a neuropsychiatric questionnaire and time-of-day (for each neuropsychiatric questionnaire, and controlling for age and gender).